<a href="https://colab.research.google.com/github/SaikumarSam/880/blob/main/Lab5/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Saikumar parre"
Registration number(s) = "2111125"
```

---

---
# Welcome to CE880
### This is your week-5 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

```python
# install sklearn package if you are using local computer and if you are using Google Coalb then skip this step
!pip install -U scikit-learn
import sys
import warnings

if not sys.warnoptions:
    warnings.simplefilter("ignore")
```

---
## Question 1: 
We’re ready to split a larger dataset to solve a regression problem. You’ll use a well-known Boston house prices dataset, which is included in sklearn. This dataset has 506 samples, 13 input variables, and the house values as the output. We have loaded dataset for you and splited into training and test set. Write code to fit linear regression on training `input (x)` and training `output (y)` and print the training score using `model.score` function. 

```python
from sklearn.datasets import load_boston
from sklearn.model_selection import train_test_split
x, y = load_boston(return_X_y=True)
x_train, x_test, y_train, y_test = train_test_split(x, y, test_size=0.4, random_state=0)
```

<!-- grade_id: my_regression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LinearRegression

def my_regression(x_train, y_train):
    """Write code to fit linear regression on training input (x) and training output (y) and 
    print the training score using model.score function. """
    # YOUR CODE HERE
    reg = LinearRegression()
    reg.fit(x_train, y_train)
    print('Regression Score',reg.score(x_train, y_train))
    # raise NotImplementedError()
    return reg.score(x_train, y_train)


my_regression(x_train, y_train)
print('Regression Accuracy:', math.isclose(my_regression(x_train, y_train), 0.7668, rel_tol = 0.05))
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Input In [1], in <cell line: 1>()
----> 1 from sklearn.linear_model import LinearRegression
      3 def my_regression(x_train, y_train):
      4     """Write code to fit linear regression on training input (x) and training output (y) and 
      5     print the training score using model.score function. """

ModuleNotFoundError: No module named 'sklearn'
```

<!-- grade_id: my_regression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_regression(x_train, y_train), 0.7668, rel_tol = 0.05)
```

---
## Question 2: 

Now, implement `RandomForestRegressor` on the same dataset and check the performance on the `test` dataset. Hint: Please read the question carefuly. 

<!-- grade_id: my_RMregression_ans, points: 0, grade: false, solution: true -->
```python
from sklearn.ensemble import RandomForestRegressor

def my_RMregression(x_train, y_train, x_test, y_test):
    """Write code to fit RandomForestRegressor on training input (x) and training output (y) and 
    print the test score using model.score function. """
    # YOUR CODE HERE
    rfr = RandomForestRegressor()
    rfr.fit(x_train, y_train)
    score = rfr.score(x_train, y_train)
    print('RandomForest Regressor Score', score)
    return score
#     raise NotImplementedError()
```

```python
# Check you solution by running this cell
import math
assert math.isclose(my_RMregression(x_train, y_train, x_test, y_test), 0.8325, rel_tol = 0.05)
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [2], in <cell line: 3>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_RMregression(x_train, y_train, x_test, y_test), 0.8325, rel_tol = 0.05)

NameError: name 'my_RMregression' is not defined
```

<!-- grade_id: my_RMregression_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_RMregression(x_train, y_train, x_test, y_test), 0.8325, rel_tol = 0.05)
```

---
## Question 3: 

Write code to fit classificaiton model LogisticRegression on training `input (x)` and training `output (y)` and 
print the predicted (x) that is your same data on which you have trained your model 

<!-- grade_id: my_classification_ans, points: 0, grade: false, solution: true -->
```python

import numpy as np
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import classification_report, confusion_matrix
x = np.arange(10).reshape(-1, 1)
y = np.array([0, 0, 0, 1, 1, 1, 1, 1, 1, 1])  # corrected input to match output, since model has already 100% accuracy

def my_classification(x, y):
    """Write code to fit classificaiton model LogisticRegression on training input (x) and training output (y) and 
    print the predicted (x) that is your same data on which you have trained your model """
    # YOUR CODE HERE
    log_reg = LogisticRegression()
    log_reg.fit(x, y)
    cm = confusion_matrix(x,y)
    print('confusion_matrix:', cm)
    score = log_reg.score(x, y)
    print('Score:', score)
    y_hat = log_reg.predict(x)
    print('y_hat:', y_hat)
    # raise NotImplementedError()
    return y_hat


```

<!-- grade_id: my_classification_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
y_hat = my_classification(x, y)
y_hat = y_hat.tolist()
assert y_hat == [0, 0, 0, 1, 1, 1, 1, 1, 1, 1]
```

```python

```