<a href="https://colab.research.google.com/github/SaikumarSam/880/blob/main/Lab7/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Saikumar parre"
Registration number(s) = "2111125"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling. Calculate the sample mean from 1000 bootstraping samples and compare with population mean. 

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=1000)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300

def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    # YOUR CODE HERE
    sample_mean = []
    for _ in range(10000):  #so B=10000
        sample_n = np.random.choice(sample, size=300)
        sample_mean.append(sample_n.mean())
    
    print(np.mean(sample_mean))
    print(pop.mean())
    print(sample.mean())
```

Output:
```
253.241
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

Output:
```
255.73952966666664
253.241
255.68
TypeError                                 Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_9144\2754433016.py in <module>
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)

TypeError: must be real number, not NoneType
```

---
## Question 2: 

#### Majority Voting
Given a list = `[100,200,300,100,200,100,100,10,30,20,20,10]`. Write a function to find the majority voting.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
a = np.array([100,200,300,100,200,100,100,10,30,20,20,10])
n = len(a)
def my_majorityvote():
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    candidate = -1
    votes = 0
     
    # Finding majority candidate
    for i in range (n):
        if (votes == 0):
            candidate = a[i]
            votes = 1
        else:
            if (a[i] == candidate):
                votes += 1
            else:
                votes -= 1
    count = 0
    for i in range (n):
        if (a[i] == candidate):
            count += 1
    
    if (count > n // 2):
        return candidate
    else:
        return -1
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 100, rel_tol = 0.05)
```

Output:
```
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_9144\529188715.py in <module>
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_majorityvote(), 100, rel_tol = 0.05)

AssertionError:
```

---
## Question 3: 

#### Decision Tree
USe the `IRIS` dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely `IRIS Setosa`, `IRIS Versicolour`, and `IRIS Virginica`, with the following attributes: `sepal length`, `sepal width`, `petal length`, and `petal width`. Your task is to train a decision tree on `training` data and return the `test` accuracy from `[0-100]`. Hint: multiply for `accuracy_score` it with 100. 

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    # YOUR CODE HERE
    clf = DecisionTreeClassifier()
    clf.fit(X_train,y_train)
    
    y_pred = clf.predict(X_test)
    score = accuracy_score(y_true = y_train, y_pred=clf.predict(X_train))
    return score
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 94.736, rel_tol = 0.05)
```

Output:
```
1.0
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp\ipykernel_9144\608619572.py in <module>
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_DTree(), 94.736, rel_tol = 0.05)

AssertionError:
```

```python

```